Identification of an early subset of cerebellar nuclei neurons in mice

Author response:

The following is the authors’ response to the previous reviews.

Reviewer #1 (Recommendations For The Authors):

(1) I was surprised to see that the Authors have failed to address my major concerns about the paper, which was in the Main text of the Review.

Previously I wrote: The major weakness of the manuscript is that it is written for a very specialized reader who has a strong background in cerebellar development, making it hard to read for eLife's general audience. It's challenging to follow the logic of some of the experiments as well as to contextualize these findings in the field of cerebellar development.

This has not been addressed. The manuscript has not been substantively changed and it is still written for a very specialized reader rather than a general reader.

We appreciate the respected reviewer’s concern and have made substantial revisions throughout the manuscript to address the points. We have simplified the technical language throughout the manuscript and included additional background information, particularly in the introduction and discussion sections, to better orient general readers. Additionally, we have clarified the logical flow of the experiments by incorporating transitional statements and summaries that explain the purpose and outcomes of each experiment (revisions are highlighted in yellow).

(2) These two have been addressed, although to be honest, I don't think that the cartoon is particularly helpful for a general audience.

Thank you for your feedback. We have replaced the cartoon with a revised version that provides more detailed information to clarify and simplify the origins of cerebellar nuclei from the caudal and rostral ends in both Atoh1+/+ and Atoh1-/- mice. We believe this will make the content more clear and informative for the general audience.

(3) My third recommendation, that they include a section in the Discussion to speculate about what these cells may become in the adult and the existence of multiple cell types with different molecular markers and projection patterns in the nuclei, has also not been addressed.

We apologize for the oversight in the previous revision. We have now added a detailed discussion in the manuscript that speculates on the potential fate of these newly identified cells in the adult cerebellum, suggesting that they may differentiate into excitatory neurons (highlighted on page 9). In addition, as noted in our previous resubmission, further direct evidence is needed from the early population of SNCA+ cells during E9 to E13. This is an ongoing focus of investigation in our lab, where we are currently using SNCA-GFP mice, part of a project for a PhD student in our lab.

Reviewer #2 (Recommendations For The Authors):

One small remaining issue: The methods text re cell counts remains confusing: n=3

EMBRYOS???

"To assess the number of OTX2-positive cells, we conducted immunohistochemistry (IHC) labeling on slides containing serial sections from embryonic days 12, 13, 14, and 15 (n=3 EMBRYOS??? at each timepoint)."

Thank you for this point and we acknowledge that, and we have revised the text in the methods section for clarity. As highlighted on page 11, “The sample size was equal to 9 embryos” and on page 16, “3 embryos were used at each time point”.

Changes from the earlier version to the revised version

Abstract

Abstract

Earlier version: Cerebellar nuclei (CN) neurons serve as the primary output of the cerebellum and originate from the cerebellar primordium at early stages of cerebellar development. Employing various methodologies, we have characterized a specific subset of CN neurons that do not originate from the rhombic lip or ventricular zone of the cerebellar primordium. Embryos were collected at early stages of development and processed for immunohistochemistry (IHC), Western blotting, in situ hybridization (ISH), embryonic culture, DiI labeling, and flow cytometry analysis (FCM). Our findings indicate that a subset of CN neurons expressing α-synuclein (SNCA), OTX2, MEIS2, and p75NTR (NGFR) are located in the rostro-ventral (rv) region of the nuclear transitory zone (NTZ), while CN neurons derived from the rhombic lip are positioned in the caudo-dorsal (cd) area of the NTZ in the cerebellar primordium. Utilizing Otx2-GFP and Atoh1−/− mice, we have determined that these cells do not originate from the germinal zone of the cerebellar primordium. These results suggest the existence of a novel extrinsic germinal zone for the cerebellar primordium, possibly the mesencephalon, from which early CN neurons originate.

Revised version: Cerebellar nuclei (CN) neurons serve as the primary output of the cerebellum and originate from the cerebellar primordium at early stages of cerebellar development. These neurons are diverse, integrating information from the cerebellar cortex and relaying it to various brain regions. Employing various methodologies, we have characterized a specific subset of CN neurons that do not originate from the rhombic lip or ventricular zone of the cerebellar primordium. Embryos were collected at early stages of development and processed for immunohistochemistry (IHC), Western blotting, in situ hybridization (ISH), embryonic culture, DiI labeling, and flow cytometry analysis (FCM). Our findings indicate that a subset of CN neurons expressing α-synuclein (SNCA), OTX2, MEIS2, and p75NTR (NGFR) are located in the rostroventral region of the nuclear transitory zone (NTZ). While CN neurons derived from the rhombic lip are positioned in the caudodorsal area of the NTZ in the cerebellar primordium. Utilizing Otx2-GFP and Atoh1−/− mice, we have determined that these cells do not originate from the germinal zone of the cerebellar primordium. These results suggest the existence of a novel extrinsic germinal zone for the cerebellar primordium, possibly the mesencephalon, from which early CN neurons originate.

Introduction

Earlier version: In the mouse, the cerebellar primordium emerges at approximately embryonic day (E) 7–E8 as a neuroepithelial swelling on the rostral lip of the fourth ventricle, which is part of the alar plate of the metencephalon (rhombomere-1) (Kebschull et al., 2023, Goldowitz and Hamre, 1998, Sotelo, 2004, Wang and Zoghbi, 2001a, Marzban et al., 2015). The cerebellar primordium is rostrally limited by the isthmus, the mesencephalon-rhombencephalon (rhombomere-1, r1) or midbrain-hindbrain boundary. Orthodenticle homeobox 2 (Otx2) and gastrulation brain homeobox 2 (Gbx2) are among the earliest genes expressed in the neuroectoderm, dividing it into anterior and posterior domains with a border that marks the isthmus. Otx2 is required for the development of the forebrain and midbrain, and Gbx2 for the anterior hindbrain (Millet et al., 1999). The isthmus, at the junction of Otx and Gbx expressing areas, is a multi-signaling center regulating genes essential for the early development of both the mesencephalon and cerebellar primordium (Joyner et al., 2000, Wurst and Bally-Cuif, 2001).

Revised version: The cerebellum, a key player in both motor and non-motor functions, is a complex structure composed of the cerebellar nuclei (CN) and cerebellar cortex. The CN, the main cerebellar output to the rest of brain, are comprised of at least three distinct neuronal types: glutamatergic, GABAergic, and glycinergic neuron (Koziol et al., 2014, Manto et al., 2012, Marzban et al., 2015). In the mouse, the cerebellar primordium, an early structure in cerebellar development, emerges at approximately embryonic day (E) 7–E8 from the alar plate of the rhombomere-1(Kebschull et al., 2023, Goldowitz and Hamre, 1998, Sotelo, 2004, Wang and Zoghbi, 2001a, Marzban et al., 2015). The cerebellar primordium is rostrally limited by the isthmus, a boundary between the mesencephalon (midbrain)-rhombencephalon (hindbrain). Two key genes, orthodenticle homeobox 2 (Otx2) and gastrulation brain homeobox 2 (Gbx2), are among the earliest genes expressed in the neuroectoderm, dividing it into anterior and posterior domains with a border that marks the isthmus. Otx2 is crucial for the development of the forebrain and midbrain, while Gbx2 is essential for the anterior hindbrain (Millet et al., 1999). The isthmus, where Otx2 and Gbx2 expression meets, is a multi-signaling center that regulates other genes essential for the early development of both the midbrain and cerebellar primordium (Joyner et al., 2000, Wurst and Bally-Cuif, 2001).

Earlier version: The cerebellar primordium contains two distinct germinal zones, which are responsible for cerebellar neurogenesis and gliogenesis; the ventrally-located ventricular zone (VZ) and dorsally-located rhombic lip (Fink et al., 2006). Early studies on the development of CN indicated that both glutamatergic and GABAergic CN neurons originate from the cerebellar VZ (Voogd and Glickstein, 1998). Further studies highlighted the involvement of the rhombic lip as an origin for glutamatergic CN neurons (between E9-E12) (Fink et al., 2006, Wang and Zoghbi, 2001a, Green and Wingate, 2014, Machold and Fishell, 2005, Kebschull et al., 2023). LIM homeobox transcription factor 1 alpha (LMX1A) regulates the development of glutamatergic CN neurons originating from the rhombic lip. It is also expressed in the nuclear transitory zone (NTZ) and is considered to be a marker for the majority of rhombic lip-derived CN neurons in addition to a subset of LMX1A positive cells that do not originate from the rhombic lip migratory stream (Chizhikov et al., 2006, Chizhikov et al., 2010).

Revised version: The cerebellar primordium contains two distinct germinal zones, which are responsible for cerebellar neurogenesis and gliogenesis; the ventrally-located ventricular zone (VZ) and dorsally-located rhombic lip (RL) (Fink et al., 2006). Early studies on the development of CN indicated that both glutamatergic and GABAergic CN neurons originate from the cerebellar VZ (Voogd and Glickstein, 1998). Further studies highlighted the involvement of the RL as an origin for glutamatergic CN neurons (between E9-E12) (Fink et al., 2006, Wang and Zoghbi, 2001a, Green and Wingate, 2014, Machold and Fishell, 2005, Kebschull et al., 2023). LIM homeobox transcription factor 1 alpha (LMX1A) regulates the development of the glutamatergic CN neurons originating from the RL. It is also expressed in the nuclear transitory zone (NTZ) and is considered to be a marker for the majority of rhombic lip-derived CN neurons in addition to a subset of LMX1A positive cells that do not originate from the RL migratory stream (Chizhikov et al., 2006, Chizhikov et al., 2010).

Earlier version: In this report, we provide evidence that a subset of CN neurons expressing OTX2, SNCA, MEIS2, and p75 Neurotrophin Receptor (p75NTR)/Nerve Growth Factor Receptor (NGFR) originate from the rostral end of the cerebellar primordium and are positioned in the rostroventral region of the NTZ during early cerebellar development. Our data suggest that the mesencephalon could represent a third germinal zone and, as such, the origin of previously unrecognized neurons that contribute to CN formation and development.

Revised version: Extensive developmental studies have shown that all cerebellar cells are produced sequentially from the VZ and the upper RL (Marzban et al., 2015). In this report, we provide evidence that a subset of CN neurons expressing OTX2, α-Synuclein (SNCA), MEIS2, and p75 neurotrophin receptor (p75NTR)/Nerve Growth Factor Receptor (NGFR) originate from the rostral end of the cerebellar primordium and are positioned in the rostroventral region of the NTZ during early cerebellar development. Our data suggest that the mesencephalon represents a third germinal zone and is the origin of previously unrecognized neurons that contribute to CN formation and development.

Results

Earlier version: In mouse embryos, the earliest neuronal populations, which are located at the rostral end of the cerebellar primordium, are immunopositive for neurofilament-associated antigens (NAA) 3A10, a marker typically expressed in neuronal somata and axons (Marzban et al., 2008, Marzban et al., 2019) (Fig. 1). To further explore these cells in the rostral end of the cerebellar primordium, we used an antibody to detect SNCA, which is expressed in CN neurons during embryogenesis (Zhong et al., 2010). In sagittal sections of the E9 cerebellar primordium, we identified SNCA+ neurons at the rostral end of the NTZ (Fig. 1E-F and Fig. Suppl. 4). To determine the location of the SNCA+ cells and compare it to that of the rhombic lip-derived LMX1A+ CN neurons (E10-12) in the NTZ, double immunofluorescence labeling of SNCA and LMX1A was performed on sagittal sections of E12 cerebellar primordium (Fig. 2 A-D, medial section and Fig. E-H, lateral section). This shows that an LMX1A+ population of CN neurons, originating from the rhombic lip, either flanks SNCA+ neurons in the medial section (Fig. 2A-D) or is located dorsally in the lateral section (Fig. 2E-H) in the NTZ. To further explore the characteristics of SNCA+ cells, double immunofluorescence labeling with OTX2 and SNCA was performed during early cerebellar development. OTX2 is highly expressed in the mesencephalon and its caudal limit coincides with the boundary of the rhombencephalon (i.e., the isthmus) (Kurokawa et al., 2004). Interestingly, results from double immunofluorescence experiments in sagittal sections of E12 cerebella showed that OTX2 was highly expressed in the mesencephalon (Fig. 3A, B), and was co-expressed with SNCA+ cells in the rostral region of the cerebellar primordium (Fig. 3 A-E). Next, we used quantitative flow cytometry analysis to determine whether OTX2+ cells co-express SNCA and/or LMX1A (Fig. 3F, G). An increase in the number of OTX2 and SNCA positive cells was observed at E14 in comparison to E12 (*p<0.05). The majority of neural cells were not SNCA+/LMX1A+, consistent with our IHC findings, or SNCA+/OTX2+, with the exception of a few cells that did show co-labeling. Of note, more SNCA+/OTX2+ than SNCA+/LMX1A+ cells were detected at E12 compared to E14.

Revised version: In mouse embryos, the earliest neuronal populations, which are located at the rostral end of the cerebellar primordium, are immunopositive for neurofilament-associated antigens (NAA; detected by 3A10 antibody), a marker typically expressed in neuronal somata and axons (Marzban et al., 2008, Marzban et al., 2019) (Fig. 1). To further explore these cells in the rostral end of the cerebellar primordium at E9, we used an antibody to detect SNCA, which is expressed in CN neurons during embryogenesis (Zhong et al., 2010). We identified SNCA+ neurons at the rostral end of the NTZ in the cerebellar primordium (Fig. 1E-F and Fig. Suppl. 4). At E12, we aimed to determine the location of the SNCA+ cells in the cerebellar primordium and compare their position with that of the rhombic lip-derived LMX1A+ CN neurons (E10-12) in the NTZ. To perform this, double immunofluorescence labeling of SNCA and LMX1A (Fig. 2 A-D, medial section and Fig. E-H, lateral section) showed that an LMX1A+ population of CN neurons, originating from the RL, either flanks SNCA+ neurons in the medial section (Fig. 2A-D) or is located dorsally in the lateral section (Fig. 2E-H) in the NTZ. This indicates that SCNA+ cells are a distinct cell population different from rhombic lip-derived LMX1A+ cells. To further explore the characteristics of SNCA+ cells, we wanted to determine whether these cells originate from the mesencephalon. OTX2 which is highly expressed in the mesencephalon, has its caudal limit at the boundary of the rhombencephalon (i.e., the isthmus)(Kurokawa et al., 2004), thus makes it an appropriate marker for identifying mesencephalon-derived cells. Interestingly, results from double immunofluorescence experiments in sagittal sections of E12 cerebella showed that OTX2 was highly expressed in the mesencephalon (Fig. 3A, B), and was co-expressed with SNCA+ cells in the rostral region of the cerebellar primordium (Fig. 3 A-E). This result suggests that SNCA+ cells originate from the mesencephalon. Next, we used quantitative flow cytometry analysis to determine whether OTX2+ cells co-express SNCA and/or LMX1A (Fig. 3F, G). An increase in the number of OTX2 and SNCA positive cells (but not LMX1A) was observed at E14 in comparison to E12 (*p<0.05). Interestingly, the majority of neural cells were not SNCA+/LMX1A+, consistent with our IHC findings, or SNCA+/OTX2+, with the exception of a few cells that did show co-labeling. Of note, more SNCA+/OTX2+ than SNCA+/LMX1A+ cells were detected at E12 compared to E14, indicating the co-expression of these markers at different time point during cerebellar development. An increase in the OTX2 expression at the rostral end of the cerebellar primordium was also confirmed by the IHC at E12-E15 (Fig. 4 A-F). The total number of OTX2+ cells in the cerebellar primordium sections were counted from E12 to E15 (Fig. 4G) and comparison between these embryonic stages indicated a slight increase in the number of OTX2+ cells between E12 and E15. Significant differences were observed between E12 and E14 (**p<0.01), E12 and E15 (***p<0.001), as well as E13 and E15 (#p<0.05).

The presence of neurons in the cerebellar primordium at E9 and spatial expression of SNCA. The cerebellar primordium immunostained with NAA and SNCA, shows that a distinct subset of neurons located at the rostral end of the developing cerebellum.

A: Dorsal view of the schematic illustration of the cerebellar primordium, mesencephalon, isthmus, and 4th ventricle. The red line indicates the sagittal plane about which the sections shown in B-D were taken.

B-D: Sagittal section through the cerebellar primordium at early E9 immunoperoxidase-labeled with NAA 3A10 shows the presence of neurons in the cerebellar primordium that cross the isthmus (i) and continue to the mesencephalon. C. A higher magnification of B. D. Differentiated neurons at E9.5 are visible; a higher magnification is shown in the inset, d.

E: Sagittal section through the cerebellar primordium at late E9. Immunofluorescence labeling of SNCA shows SNCA+ (green) expressing CN neurons in the mesencephalon, at the isthmus (i) and in the rostral part of the cerebellar primordium (cb). F. A higher magnification of E.

Abbreviations: 4thv, 4th ventricle; cb, cerebellum; c, caudal; d, dorsal; i, isthmus; m, mesencephalon; r, rostral; v, ventral

Scale bar = 100 µm in B; 50 µm in C and E; 20 µm in D

SNCA expression relative to Lmx1a in the early developing mouse cerebellum. The sagittal section through the cerebellar primordium at E12.5 reveals double immunofluorescence labeling with SNCA and LMX1A antibodies in both medial (panels A-D) and lateral (panels E-H) sections.

A–G: SNCA (green, A and E) and LMX1A (red, B and F) immunopositive cells are located at the NTZ. SNCA+ cells continue to the mesencephalon, and LMX1A+ cells toward the rhombic lip. Merged images show that the SNCA+ cells form a population of CN neurons distinct from the rhombic lip-derived cells (LMX1A+) in the NTZ (C and G). D and H show a higher magnification of C and G, respectively.

Abbreviations: 4thV, 4th ventricle; cb, cerebellum; c, caudal; d, dorsal; i, isthmus; m, mesencephalon; rl, rhombic lip; NTZ, nuclear transitory zone; r, rostral; v, ventral.

Scale bar = 100 µm.

Expression of SNCA and OTX2 in the developing cerebellar primordium. Double immunofluorescence labeling with SNCA and OTX2 antibodies in sagittal sections of the cerebellar primordium at E12.5.

A–E: Double immunolabeling of SNCA (green) and OTX2 (red) of a sagittal section of the cerebellar primordium at E12.5 A. High OTX2 immunoreactivity was detected in the mesencephalon. SNCA+ cells in the mesencephalon accompany OTX2+ cells across the isthmus (i) and in the NTZ. B. higher magnification of A. C-E. Immunolabeling with SNCA (green, C) and OTX2 (red, D) and merged (E) shows co-expression in the NTZ.

F-G: Cerebellar cells at E12 and E14 were dissociated and immunolabeled for SNCA + OTX2, SNCA + LMX1A, or each antibody individually (F). Quantitative analysis was performed by flow cytometry to detect SNCA+, OTX2+, LMX1A+, SNCA+/ LMX1A+ and SNCA+/ OTX2+ cells (G). No significant differences were observed for the number of LMX1A+ cells between E12 and E14. An increase in SNCA and OTX2 positive cells was shown at E14 in comparison to E12 (*P<0.05). The number of cells positive for SNCA/LMX1A or SNCA/OTX2 was very low at E12 and E14.

Abbreviations: 4thv, 4th ventricle; cb, cerebellum; c, caudal; d, dorsal; i, isthmus; m, mesencephalon; rl, rhombic lip; NTZ, nuclear transitory zone; r, rostral; v, ventral.

Scale bar = 200 µm in A; 50 µm in B; 20 µm in C, D, and E.

Removed in the revised version: IHC performed on sagittal sections of E12 to E15 cerebella showed that OTX2+ cells are present at the rostral end of the cerebellar primordium (Fig. 4 A-F). To evaluate the number of OTX2+ cells in the developing rostral cerebellum, the total number of OTX2+ cells in the cerebellar primordium sections were counted from E12 to E15 (Fig. 4G). A comparison between these embryonic stages indicated a slight increase in the number of OTX2+ cells between E12 and E15. Significant differences were observed between E12 and E14 (**p<0.01), E12 and E15 (***p<0.001), as well as E13 and E15 (#p<0.05).

The spatial and temporal expression of OTX2 in the developing cerebellum of mice. Sagittal sections through the cerebellar primordia at E12, 13, 14, and 15 were analyzed for peroxidase immunoreactivity of OTX2.

A–F: Sagittal section through cerebellar primordium at E12 (A; higher magnification shown in B), E13 (C; higher magnification shown in D), E14 (E), and E15 (F). High OTX2 immunoreactivity at the mesencephalon is evident, and a few OTX2+ cells cross the isthmus and position at the rostral part of the cerebellar primordium at the NTZ. Examples of Otx2 expressing cells are indicated by arrowheads in panels B, D, E, and F. In panels E and F, the isthmus and adjacent cerebellar area are showed in the same orientation as panels A-D (I; isthmus). The orientation label (r-c and d-v; rostral caudal and dorsal ventral) in panel A applies to all panels.

Earlier version: G: Comparison of the number of OTX2 positive cells in the cerebellar primordium from E12 to E15. Results indicate a slight increase in the number of OTX2 positive cells over time. Significant differences were observed between E12 and E14 (**P<0.01), E12 and E15 (***P<0.001), as well as E13 and E15 (#P<0.05). Data were analyzed by One-Way ANOVA followed by a Tukey’s multiple comparison test.

Revised version: G: Comparison of the number of OTX2 positive cells in the cerebellar primordium from E12 to E15. Results indicate a slight increase in the number of OTX2 positive cells over time. Significant differences were observed between E12 and E14 (**P<0.01), E12 and E15 (***P<0.001), as well as E13 and E15 (#P<0.05). Data were analyzed by One-Way ANOVA followed by a Tukey’s multiple comparison test. Abbreviations: 4thv, 4th ventricle; cb, cerebellum; c, caudal; d, dorsal; i, isthmus; m, mesencephalon; rl, rhombic lip; r, rostral; v, ventral.

Removed in the revised version: Abbreviations: 4thv, 4th ventricle; cb, cerebellum; c, caudal; d, dorsal; i, isthmus; m, mesencephalon; rl, rhombic lip; r, rostral; v, ventral.

Scale bar = 200 µm in A and C; 50 µm in F (applies to B, D, E, and F).

Earlier version: To confirm the presence of OTX2+ cells in the rostral end of the cerebellar primordium, we used Otx2-GFP transgenic mice and performed double immunofluorescence labeling with OTX2 and GFP on E13 sagittal sections (Fig. 5). Results showed that an extension of OTX2-GFP positive cells, highly prevalent in the mesencephalon, crossed the isthmus and terminated in the rostral end of the cerebellar primordium in the NTZ (Fig. 5B, C, and D). To further validate our observations, we employed RNAscope fluorescence in situ hybridization (FISH) at E12 to detect the presence of Otx2 mRNA in the cerebellar primordium (Fig. 5 E-J). Otx2 mRNA was highly expressed in the mesencephalon and caudally extended to the rostral cerebellar primordium in the NTZ (Fig. 5 E-J and Suppl. Fig. 4).

Revised version: To further confirm the presence of OTX2+ cells in the rostral end of the cerebellar primordium, we used Otx2-GFP transgenic mice and performed double immunofluorescence labeling with OTX2 and GFP on E13 sections (Fig. 5). Results showed that an extension of OTX2-GFP positive cells, highly prevalent in the mesencephalon, crossed the isthmus and terminated in the rostral end of the cerebellar primordium in the NTZ (Fig. 5B, C, and D). To further validate our observations, we employed RNAscope fluorescence in situ hybridization (FISH) at E12 to detect the presence of Otx2 mRNA in the cerebellar primordium (Fig. 5 E-J). Otx2 mRNA was highly expressed in the mesencephalon and caudally extended to the rostral cerebellar primordium in the NTZ (Fig. 5 E-J and Suppl. Fig. 4).

Otx2 expression in developing cerebella of Otx2-GFP transgenic and control mice. Sagittal sections through the cerebellar primordium of Otx2-GFP transgenic mice at E13 are shown in panels A-D, and RNAscope fluorescence in situ hybridization (FISH) of Otx2 at E12 in wild-type mice in panels E-J.

A-D: DAPI labels the outline of the cerebellar primordium and mesencephalon (A, blue). GFP expression, which is enhanced by immunofluorescence labeling with anti-GFP (B, green), and immunofluorescence labeling for OTX2 (C, red), reveal co-labeled cells (D, merged) in the mesencephalon and NTZ at the rostral end of the cerebellar primordium. Arrows indicate the isthmus.

E-G: Merged channels of the in situ hybridization of Otx2 mRNA probe counterstained with DAPI at low (E) and high (F and G) magnification captured by confocal microscopy.

H-J: The Otx2 mRNA signal was strong in the mesencephalon and extended as a tail through the isthmus to the rostral cerebellar primordium in the NTZ. The isthmus is indicated by arrows in (E, F, H, I) and a line in (J).

Abbreviations: cb, cerebellum; c, caudal; d, dorsal; m, mesencephalon; rl, rhombic lip; ntz, nuclear transitory zone; r, rostral; v, ventral.

Scale bar = 200 µm in D (applies to A-D); 100 µm in H (applies to E and H); 20 µm in J (applies to F - J).

Earlier version: To evaluate whether OTX2+ cells are present in the absence of Snca gene expression, we used Pap−/−; Snca−/−mice. OTX2 immunoperoxidase staining showed that even in the absence of Snca expression, OTX2+ cells still terminated in the rostral cerebellar primordium (Suppl. Fig. 1). Nerve growth factor receptor (p75NTR; Ngfr), which regulates neuronal proliferation and differentiation (Jiang et al., 2008, Bernabeu and Longo, 2010, Dechant and Barde, 2002), has also been shown to be expressed in the rostral end of the cerebellar primordium in Pap−/−; Snca−/− mice (Rahimi-Balaei et al., 2019b). To determine whether SNCA+ neurons are p75NTR immunopositive, double immunolabeling with SNCA and p75NTR was performed. Section IHC revealed that SNCA+ cells in the NTZ exhibited cell membrane expression of p75NTR (Fig. 6A–D). This data was further confirmed by Western blot analysis of SNCA and p75NTR protein expression in embryonic cerebellar lysates (E11, 13 and 15; Fig. 6d). In addition, to conclusively determine that p75NTR is expressed in the membranes of SNCA+ cells, we utilized primary dissociated cultures of cerebellum at E10, DIV 4 (Fig. 6E-H). IHC demonstrated that p75NTR was indeed expressed in the membranes of SNCA+ cells (with punctate appearance, Fig. 6E–H (Dechant and Barde, 2002)).

Revised version: To evaluate whether OTX2+ cells are present in the absence of Snca gene expression, we used Pap−/−; Snca−/−mice. OTX2 immunoperoxidase staining showed that even in the absence of Snca expression, OTX2+ cells still terminated in the rostral cerebellar primordium (Suppl. Fig. 1). This indicates that the expression of Otx2 in these cells is independent of SNCA. We previously reported that nerve growth factor receptor (p75NTR; Ngfr), a neuronal proliferation and differentiation marker, was expressed in the rostral end of the cerebellar primordium in Pap−/−; Snca−/− mice (Jiang et al., 2008, Bernabeu and Longo, 2010, Dechant and Barde, 2002, Rahimi-Balaei et al., 2019b). To determine whether SNCA+ neurons are p75NTR immunopositive, double immunolabeling with SNCA and p75NTR revealed that SNCA+ cells in the NTZ exhibited cell membrane expression of p75NTR (Fig. 6A–D). This data was further confirmed by Western blot analysis of SNCA and p75NTR protein expression in embryonic cerebellar lysates (E11, 13 and 15; Fig. 6d). In addition, to conclusively determine that p75NTR is expressed in the membranes of SNCA+ cells, we utilized primary dissociated cultures of cerebellum at E10, DIV 4 (Fig. 6E-H). IHC demonstrated that p75NTR was indeed expressed in the membranes of SNCA+ cells (with punctate appearance, Fig. 6E–H (Dechant and Barde, 2002)).

SNCA immunopositive cells express p75NTR (NGFR) in developing cerebellum in vivo and in vitro.

Earlier version: A–D: Double immunofluorescence labeling with SNCA (A, green) and p75NTR (B, red) antibodies reveals co-labeled cells (C, merged) in the NTZ; a higher magnification of the NTZ is shown in D. d. Western blot analysis of SNCA and p75NTR expression during cerebellar development. Immunoblots of total cerebellar lysates from embryos at E11, E13 and E15 indicate an increase in expression of SNCA and p75NTR from E11 to E15. Equal protein loading was confirmed by β-actin expression.

Revised version: A–D: Double immunofluorescence labeling with SNCA (A, green) and p75NTR (B, red) antibodies reveals co-labeled cells (C, merged) in the NTZ; a higher magnification of the NTZ is shown in D.

Added in the revised version: d. Western blot analysis of SNCA and p75NTR expression during cerebellar development. Immunoblots of total cerebellar lysates from embryos at E11, E13 and E15 indicate an increase in expression of SNCA and p75NTR from E11 to E15. Equal protein loading was confirmed by β-actin expression.

Earlier version: E–H: Primary dissociated cultures of cerebellum obtained from an E10 mouse embryo (DIV 4), double immunofluorescence stained for SNCA (E: green) and p75NTR (F: red) (merged image shown in G). H is a higher magnification of G; punctuate cellular p75NTR immunoreactivity is marked by arrow heads.

Revised version: E–H: Primary dissociated cultures of cerebellum obtained from an E10 mouse embryo (DIV 4), double immunofluorescence stained for SNCA (E: green) and p75NTR (F: red) (merged image shown in G). H is a higher magnification of G; punctuate cellular p75NTR immunoreactivity is marked by arrow heads. Abbreviations: cb, cerebellum; c, caudal; d, dorsal; NTZ, nuclear transitory zone; r, rostral; v, ventral.

Removed in the revised version: Abbreviations: cb, cerebellum; c, caudal; d, dorsal; NTZ, nuclear transitory zone; r, rostral; v, ventral.

Scale bar = 50 µm in A–C and D–F; 20 µm in H; 10 µm in G.

Earlier version: To determine whether OTX2 and SNCA immunopositive cells exist in the absence of rhombic lip-derived CN neurons, we used Atoh1−/− mice (Fig. 7 A-C). Immunostaining revealed the presence of SNCA+ and OTX2+ cells at the rostral end of the cerebellar primordium in these animals. To validate our findings, we evaluated MEIS2, a transcription factor typically expressed in the NTZ, in Atoh1+/+ and Atoh1−/− sagittal sections of the cerebellar primordium at E12. Notably, our results show the presence of two distinct sets of MEIS2+ cells in the NTZ of Atoh1+/+ embryos: rhombic lip-derived CN neurons located in the caudodorsal region, which do not express SNCA (MEIS2+/SNCA−), and a subset of MEIS2+/SNCA+ CN neurons situated in the rostroventral region of the NTZ (Fig. 7D-I). In Atoh1−/− sections, MEIS2+/SNCA+ cells were still present in the rostroventral region of the NTZ, whereas the subpopulation of MEIS2+/SNCA− rhombic lip-derived CN neurons was absent in the caudodorsal region of the NTZ (Fig. 7 J-O).

Revised version: Atonal BHLH Transcription Factor 1 (ATOH1) plays a crucial role in regulating the development of glutamatergic CN neurons derived from the RL at E12, prior to the formation of other rhombic lip-derived cells (granule cells and unipolar brush cells) (Ben-Arie et al., 1997, Wang et al., 2005). To determine whether OTX2 and SNCA immunopositive cells exist in the absence of rhombic lip-derived CN neurons, we used Atoh1−/− mice (Fig. 7 A-C). Immunostaining revealed the presence of SNCA+ and OTX2+ cells at the rostral end of the cerebellar primordium at E12, which reveals this cell population has no RL origin. To validate our findings, we evaluated MEIS2, a transcription factor typically expressed in the NTZ and glutamatergic CN neurons (Willett et al., 2019), in Atoh1+/+ and Atoh1−/− cerebellar sections at E12. Notably, our results showed the presence of two distinct sets of MEIS2+ cells in the NTZ of Atoh1+/+ embryos: 1) rhombic lip-derived CN neurons located in the caudodorsal region, which do not express SNCA (MEIS2+/SNCA-), and 2) a subset of MEIS2+/SNCA+ CN neurons situated in the rostroventral region of the NTZ (Fig. 7D-I). In Atoh1−/− sections, MEIS2+/SNCA+ cells were still present in the rostroventral region of the NTZ, whereas the subpopulation of MEIS2+/SNCA- rhombic lip-derived CN neurons was absent in the caudodorsal region of the NTZ (Fig. 7 J-O).

SNCA, OTX2, and MEIS2 expression in the developing cerebella of Atoh1 null and control mice. Sagittal sections through the cerebellar primordia of Atoh1+/+ and Atoh1−/− embryos at E12.5; double immunofluorescence labeling with SNCA + OTX2 and SNCA + MEIS2 antibodies is demonstrated.

A–C: SNCA (green, A) and OTX2 (red, B) immunopositive cells are present in the mesencephalon and NTZ of Atoh1 KO mouse (merged channels, C).

D-F: In E12.5 Atoh1+/+ sagittal sections, SNCA+ (green, D) cells are observed in the mesencephalon and NTZ. MEIS2+ (red, E) cells are present in the mesencephalon and NTZ in addition to another population of MEIS2+ cells in the dorsal region of the NTZ which extends to the rhombic lip. Merged image (F) confirms the presence of two distinct sets of MEIS2+ cells in the NTZ: rhombic lip-derived CN neurons located in the caudodorsal (cd) region, which do not express SNCA (MEIS2+/SNCA-), and a subset of MEIS2+/SNCA+ CN neurons situated in the rostroventral (rv) region of the NTZ.

G-I: Higher magnification of D-F.

J-L: In E12.5 Atoh1−/− sagittal sections, the expression of SNCA (green, J) shows no change as compared to Atoh1+/+, whereas Meis2 abundance (red, K) exhibited a different pattern. Thus, MEIS2+/SNCA+ cells were still present in the rostroventral region of the NTZ. However, the subpopulation of MEIS2+/SNCA- rhombic lip-derived CN neurons were absent in the caudodorsal region of the NTZ.

M-O: Higher magnification of J-I.

Abbreviations: 4thV, 4th ventricle; cb, cerebellum; cd, caudodorsal; i, isthmus; m, mesencephalon; rl, rhombic lip; NTZ, nuclear transitory zone; rv, rostroventral.

Scale bar = 100 µm

To explore if the rostroventral subset of CN neurons derived from an external cerebellar germinal zone, possibly the mesencephalon, we utilized FAST DiI as a neuronal tracer to mark cells within a potential source region. FAST DiI was applied to the dorsum of the caudal mesencephalon at E9, where it was maintained for 4 days (Fig. 8 A, a). At DIV 4, cells stained with DiI in the dorsum of the mesencephalon were present in both rostral (Fig. 8B) and caudal (Fig. 8C) directions. Sections from the cerebellar primordium revealed the presence of DiI-labeled cells within the cultured embryo (Fig. Suppl. 2). To investigate the earliest DiI-positive cell population in the mesencephalon and avoid unwanted cell staining due to long DiI exposure, we focused on early cerebellar development at E9; it has been reported that one of the premier neuronal populations in the CNS is present at this time in the mesencephalon and projects caudally (Stainier and Gilbert, 1990, Easter et al., 1993). To determine if an early generation of mesencephalic DiI-labeled cells is present within the cerebellar primordium, we limited cellular DiI exposure to only 24 hours (Fig. 8 D, d), after which it was removed (Fig. Suppl. 3). Intriguingly, almost all DiI+ cells migrated caudally and presented in the cerebellar primordium (Fig. 8E, e and Fig. Suppl. 3). DiI staining was clear and strong in cells as shown in Fig. 8 E, e and Fig. Suppl. 3 (arrowhead); it should be noted, however, that some DiI staining was evident in other cells in the rostral end of the cerebellar primordium, although the signal appeared relatively weak (Suppl. Fig. 3).

Labeling of mesencephalic cells during early development in mice using embryonic culture. Fast DiI was applied to an embryo at E9 and kept for 4 days (DIV 4, panels A-C) or removed after 24 hours (DIV 1, panels D-E).

A, a: Fast DiI was inserted in the mesencephalon at E9 (DIV 0), arrowhead shows insertion location of DiI crystal in the mesencephalon. The arrow indicates the isthmus. Inset “a” shows the whole embryo and indicates the site of DiI insertion.

B-C: 4 days post DiI insertion, DiI positive cells directed both rostrally to the mesencephalon (B) and caudally to the rostral cerebellar primordium (C).

D, d: Fast DiI was inserted in the mesencephalon at E9 (DIV 0) (indicated by arrowhead) and removed at DIV1, arrow shows the isthmus. Inset “d” shows the whole embryo and indicates the site of DiI insertion. E, e: DiI positive cells present in cerebellar primordium at DIV 6 after removal of Fast DiI at DIV1.

Earlier version: Abbreviations: cb, cerebellum; c, caudal; d, dorsal; m, mesencephalon; r, rostral; v, ventral.

Revised version: Abbreviations: cb, cerebellum; c, caudal; d, dorsal; m, mesencephalon; r, rostral; v, ventral. Scale bar = 500 µm in a, d; 200 µm in A, D; 100 µm in B-C, E.

Removed in the revised version: Scale bar = 500 µm in a, d; 200 µm in A, D; 100 µm in B-C, E.

Earlier version: A schematic illustration of the developing cerebellum at ∼E12, created by BioRender.com. The expression patterns of SNCA, MEIS2, and OTX2 within the cerebellar primordium indicate a subset of CN neurons with origins distinct from the rhombic lip (rl).

Revised version: A schematic illustration of the developing cerebellum at ~E12, created using BioRender.com. The expression patterns of SNCA, MEIS2, and OTX2 within the cerebellar primordium indicate a subset of CN neurons with origins distinct from the rhombic lip (rl).

Abbreviations: c, caudal; d, dorsal; m, mesencephalon; ntz, nuclear transitory zone; r, rostral; rl, rhombic lip; v, ventral.

Discussion

Earlier version: Previously, it was commonly held that all CN neurons had a singular origin in the ventricular zone (Pierce, 1975, Miale and Sidman, 1961). However, recent discoveries have challenged this notion, revealing a dual origin for CN neurons. Some originate from the ventricular zone, while others arise from the rhombic lip (Wang and Zoghbi, 2001b, Wang et al., 2005, Ben-Arie et al., 1997). Genetic fate mapping further suggested that most glutamatergic CN projection neurons may arise from the rhombic lip (Ben-Arie et al., 1997, Machold and Fishell, 2005, Fink et al., 2006). Transcription factor expression patterns indicate that CN neurons migrate from the rhombic lip to the NTZ through a subpial stream pathway, while sequentially expressing the genes Lmx1a, Pax6, Tbr2, and Tbr1 (Fink et al., 2006).

Revised version: Previously, it was commonly held that all CN neurons had a singular origin in the VZ (Pierce, 1975, Miale and Sidman, 1961). However, recent discoveries have challenged this notion, revealing a dual origin for CN neurons. Some originate from the VZ, while others arise from the RL (Wang and Zoghbi, 2001b, Wang et al., 2005, Ben-Arie et al., 1997). Genetic fate mapping further suggested that glutamatergic CN projection neurons may arise from the RL (Ben-Arie et al., 1997, Machold and Fishell, 2005, Fink et al., 2006). Transcription factor expression patterns indicate that CN neurons migrate from the RL to the NTZ through a subpial stream pathway, while sequentially expressing the genes Lmx1a, Pax6, Tbr2, and Tbr1 (Fink et al., 2006).

Earlier version: In this study, we characterized a subset of CN neurons that do not originate from the cerebellar primordium germinal zones during early cerebellar development. We demonstrated a newly identified subset of CN neurons, which are SNCA+/ OTX2+/ MEIS2+/ p75NTR+/ LMX1A−, in the rostroventral region of the NTZ in the rostral end of the cerebellar primordium. This suggests the existence of a new germinal zone during cerebellar neurogenesis.

Revised version: In this study, we characterized a subset of CN neurons that do not originate from the cerebellar primordium germinal zones (RL and VZ) during early cerebellar development. We demonstrated a newly identified subset of CN neurons, which are SNCA+/ OTX2+/ MEIS2+/ p75NTR+/ LMX1A-, in the rostroventral region of the NTZ of the cerebellar primordium. This suggests the existence of a new germinal zone during cerebellar neurogenesis.

Earlier version: The exact origin of the SNCA+/ OTX2+/ MEIS2+/ p75NTR+/ LMX1A− CN neurons is currently unclear; however, they do not originate from the rhombic lip or ventricular zones. Rhombic lip-derived Tbr1+/Lmx1a+ CN neurons are born at E9, but do not reach the NTZ until ∼E11 (Kebschull et al., 2023, Machold and Fishell, 2005, Wang et al., 2005, Manto et al., 2012, Fink et al., 2006, Rahimi-Balaei et al., 2018). Our results revealed that SNCA+ cells are a group of differentiating neurons (NAA 3A10+) present in the NTZ at E9, before the arrival of any neurons that originate from the rhombic lip. The majority of SNCA+ CN neurons are not LMX1A+; however, a small proportion of cells seems to exhibit co-expression. This suggests that in the early stages of CN neurogenesis, the pattern of protein expression in SNCA+ neurons may be changing, possibly with respective down- and upregulation of Snca and Lmx1a+. Conversely, this may not occur at all, and ‘co-labeling’ could have been observed due to overlapping cells. A similar display of SNCA expression has recently been demonstrated in oligodendrocyte development and maturation (Djelloul et al., 2015). Our findings indicated that CN neuron somata are highly positive for SNCA until E14, after which SNCA expression is refined to PCs at P0 (Zhong et al., 2010) and axonal projections, indicating that SNCA+ cells at the rostral end of the cerebellar primordium are associated with CN establishment in the NTZ. Our results from studies using Pap−/− ; Snca−/− mice showed that lack of SNCA during cerebellar development did not affect the rostroventral subset of CN neurons (Rahimi-Balaei et al., 2019a), which still expressed OTX2 and terminated in the rostral cerebellar primordium. In adult mice, postnatal expression of SNCA has been reported in medial CN neurons (Fujita et al., 2020).

Revised version: The exact origin of the SNCA+/ OTX2+/ MEIS2+/ p75NTR+/ LMX1A- CN neurons is currently unclear; however, they do not originate from the RL or VZ. Rhombic lip-derived TBR1+/LMX1A+ CN neurons are born at E9, but do not reach the NTZ until ~E11 (Kebschull et al., 2023, Machold and Fishell, 2005, Wang et al., 2005, Manto et al., 2012, Fink et al., 2006, Rahimi-Balaei et al., 2018). Our results revealed that SNCA+ cells are a group of differentiating neurons (NAA 3A10+) present in the NTZ at E9, before the arrival of any neurons that originate from the RL. The majority of SNCA+ CN neurons are not LMX1A+; however, a small proportion of cells seems to exhibit co-expression. This suggests that in the early stages of CN neurogenesis, the pattern of protein expression in SNCA+ neurons may be changing, possibly with respective down- and upregulation of Snca and Lmx1a. Conversely, this may not occur at all, and ‘co-labeling’ could have been observed due to overlapping cells. A similar display of SNCA expression has recently been demonstrated in oligodendrocyte development and maturation (Djelloul et al., 2015). Our findings indicated that CN neuron somata are highly positive for SNCA until E14, after which SNCA expression is refined to Purkinje cells at P0 (Zhong et al., 2010) and axonal projections, suggesting that SNCA+ cells at the rostral end of the cerebellar primordium are associated with CN establishment in the NTZ. Our results from studies using Pap−/−; Snca−/− mice showed that lack of SNCA during cerebellar development did not affect the formation of rostroventral subset of CN neurons (Rahimi-Balaei et al., 2019a), which still expressed OTX2 and terminated in the NTZ. While our comprehensive fate-mapping study of SNCA-expressing CN neurons is still in progress, previous research has documented SNCA expression in the medial CN neurons of adult mice. (Fujita et al., 2020). Excitatory CN neurons were previously classified into two main types. However, recent research has further subdivided these neurons into spatially distinct subpopulations within each cerebellar nucleus. (Kebschull et al., 2023, Krishnamurthy et al., 2024). Specifically, excitatory neurons in the medial CN have been divided into four subpopulations based on the expression of three marker proteins, including SNCA (Fujita et al., 2020). Notably, SNCA+ neurons are predominantly located ventrally in the caudal medial CN. Further studies using mice with mutations in En1 and En2, which are crucial for regulating the function of the isthmic organizer, have demonstrated a loss of a subpopulation of excitatory neurons in the medial CN (Krishnamurthy et al., 2024). This finding suggests that the early SNCA+ subpopulation may consist of excitatory neurons that did not originate from the RL.

Earlier version: Furthermore, our observations reveal the expression of MEIS2 in the NTZ (Fig. 7 and (Morales and Hatten, 2006)), along with several other NTZ markers such as POU3F1, BRN2, IRX3, and MEIS1 as documented in (Kebschull et al., 2023, Wu et al., 2022), c-Ret and Tlx3 (Fig. Suppl. 4), indicating the existence of discrete subpopulations of precursors within the CN. MEIS2 (homeobox transcription regulator) is expressed in the mesencephalic alar plate of both mice (Cecconi et al., 1997) and chicks (Agoston and Schulte, 2009, Bobak et al., 2009). MEIS2 plays a crucial role in the development of cranial nerves and craniofacial structures (Machon et al., 2015). It has been suggested that mesencephalon development is regulated by Meis2 cross-talk with Otx2 (Agoston and Schulte, 2009). Notably, the expression of MEIS2 in CN neurons is remarkably distinct, allowing for a clear demarcation of the caudodorsal (SNCA−/MEIS2+) and rostroventral (SNCA+/MEIS2+) regions of the NTZ, whether originating from the rhombic lip or not (Fig. 7D-I). CN neurons derived from the rhombic lip require ATOH1 expression, as they are absent in an Atoh1-null mutant (Wang et al., 2005, Machold and Fishell, 2005, Zordan et al., 2008, Florio et al., 2012, Ben-Arie et al., 1997). In the Atoh1-null cerebellar primordium, the caudodorsal (SNCA−/MEIS2+) subset of CN is absent in the NTZ, while the rostroventral (SNCA+/MEIS2+) population is still present, indicating that the Meis2+ rostroventral subset of the CN in the NTZ develops independently from Atoh1.

Revised version: Furthermore, our observations reveal the expression of MEIS2 in the NTZ (Fig. 7 and (Morales and Hatten, 2006)). The expression of MEIS2 in NTZ along with several other NTZ markers such as POU3F1, BRN2, IRX3, and MEIS1 as documented in (Kebschull et al., 2023, Wu et al., 2022), c-Ret and Tlx3 (Fig. Suppl. 4), indicating the existence of discrete subpopulations of precursors within the CN. MEIS2 (homeobox transcription regulator) is expressed in the mesencephalic alar plate of both mice (Cecconi et al., 1997) and chicks (Agoston and Schulte, 2009, Bobak et al., 2009). MEIS2 plays a crucial role in the development of cranial nerves and craniofacial structures (Machon et al., 2015). It has been suggested that mesencephalon development is regulated by Meis2 cross-talk with Otx2 (Agoston and Schulte, 2009). Notably, the expression of MEIS2 in CN neurons is remarkably distinct, allowing for a clear demarcation of the caudodorsal (SNCA-/MEIS2+) and rostroventral (SNCA+/MEIS2+) regions of the NTZ, whether originating from the RL or not (Fig. 7D-I). CN neurons derived from the RL require ATOH1 expression, as they are absent in an Atoh1-null mutant (Wang et al., 2005, Machold and Fishell, 2005, Zordan et al., 2008, Florio et al., 2012, Ben-Arie et al., 1997). In the Atoh1-null cerebellar primordium, the caudodorsal (SNCA-/MEIS2+) subset of CN is absent in the NTZ, while the rostroventral (SNCA+/MEIS2+) population is still present, indicating that the Meis2+ rostroventral subset of the CN in the NTZ develops independently from Atoh1.

Earlier version: What are the potential origins of SNCA+/ OTX2+/ MEIS2+/ p75NTR+/ LMX1A− cells in the NTZ during the early stages of cerebellar development? Our study suggests that the origin of these cells may be guided by a continuous flow of rostroventral cells toward the rostral end (mesencephalon), and caudodorsal cells toward the caudal end (rhombic lip) of the cerebellar primordium. It is well established that Otx2 is required for the development of the fore- and midbrain, while Gbx2 is necessary for anterior hindbrain development (Simeone et al., 1992, Alvarado-Mallart, 2005, Li and Joyner, 2001, Joyner et al., 2000). The expression of OTX2 in the rostral and GBX2 in the caudal neural tube, is considered limited to the mesencephalon/metencephalon boundary (isthmus) (Alvarado-Mallart, 2005). However, a study by Martinez et al. on chick/quail chimeras determined that the rostral portion of the cerebellar primordium is located more rostrally in the so-called ‘mesencephalic’ alar plate (Martinez and Alvarado-Mallart, 1989). Recently, Isl1 positive cells in the anterior cerebellum were shown to exhibit residual OTX2 protein activity (Wizeman et al., 2019). Several other studies have shown that rostral cerebellar development is associated with factors expressed in the caudal mesencephalon, such as engrailed family En1 and −2 (Joyner et al., 1991, Hanks et al., 1995, Millen et al., 1995), and Acp2 (Bailey et al., 2013, Bailey et al., 2014), which are regulated by the multi-signaling center known as the isthmic organizer (Martinez et al., 2013). Furthermore, by employing embryonic cultures in conjunction with DiI labeling, we determined that the source of the rostroventral subset of CN neurons appears to be an external cerebellar germinal zone, possibly originating from the mesencephalon, which is in line with other studies (Nichols and Bruce, 2006, Wizeman et al., 2019). While our current study hints at the possibility of the caudal mesencephalon serving as a novel extrinsic germinal zone for the cerebellar primordium, further investigation is required. Genetic inducible fate mapping and long-term follow-up studies are essential to gain a comprehensive understanding of the origin, fate, and connectivity of, as well as the role played by the SNCA+/ OTX2+/ MEIS2+/ p75NTR+/ LMX1A− rostroventral subset of CN neurons in the development of the cerebellum.

Revised version: What are the potential origins of SNCA+/ OTX2+/ MEIS2+/ p75NTR+/ LMX1A- cells in the NTZ during the early stages of cerebellar development? Our study suggests that the origin of these cells may be guided by a continuous flow of rostroventral cells toward the rostral end (mesencephalon), and caudodorsal cells toward the caudal end (RL) of the cerebellar primordium. It is well established that Otx2 is required for the development of the fore- and midbrain, while Gbx2 is necessary for anterior hindbrain development (Simeone et al., 1992, Alvarado-Mallart, 2005, Li and Joyner, 2001, Joyner et al., 2000). The expression of OTX2 in the rostral and GBX2 in the caudal neural tube, is considered limited to the isthmus (Alvarado-Mallart, 2005). However, a study by Martinez et al. on chick/quail chimeras determined that the rostral portion of the cerebellar primordium is located more rostrally in the so-called ‘mesencephalic’ alar plate (Martinez and Alvarado-Mallart, 1989). Recently, Isl1 positive cells in the anterior cerebellum were shown to exhibit residual OTX2 protein activity (Wizeman et al., 2019). Several other studies have shown that rostral cerebellar development is associated with factors expressed in the caudal mesencephalon, such as engrailed family En1 and −2 (Joyner et al., 1991, Hanks et al., 1995, Millen et al., 1995), and Acp2 (Bailey et al., 2013, Bailey et al., 2014), which are regulated by the multi-signaling center known as the isthmic organizer (Martinez et al., 2013). Furthermore, by employing embryonic cultures in conjunction with DiI labeling, we determined that the source of the rostroventral subset of CN neurons appears to be an external cerebellar germinal zone, possibly originating from the mesencephalon, which is in line with other studies (Nichols and Bruce, 2006, Wizeman et al., 2019). While our current study hints at the possibility of the caudal mesencephalon serving as a novel extrinsic germinal zone for the cerebellar primordium, further investigation is required. Genetic inducible fate mapping and long-term follow-up studies are essential to gain a comprehensive understanding of the origin, fate, and connectivity of, as well as the role played by the SNCA+/ OTX2+/ MEIS2+/ p75NTR+/ LMX1A- rostroventral subset of CN neurons in the development of the cerebellum.

In conclusion, our study indicates that the SNCA+/ OTX2+/ MEIS2+/ p75NTR+/ LMX1A- rostroventral subset of CN neurons do not originate from the well-known distinct germinative zones of the cerebellar primordium. Instead, our findings suggest the existence of a previously unidentified extrinsic germinal zone, potentially the mesencephalon.

Material and methods

Animal maintenance

Earlier version: All animal procedures were performed in accordance with institutional regulations and the Guide to the Care and Use of Experimental Animals from the Canadian Council for Animal Care. For this study, we used embryos from 47 CD1 timed-pregnant mice at embryonic day (E) 9 to E18 (total number of embryos was 409). Our approach to sample size was to include enough samples to reach a sufficient level, which was determined empirically (in most of the experiments, n was equal to 9).Timed-pregnant, prostatic acid phosphatase (PAP) mutant (Pap−/−) (Zylka et al., 2008, Hokin and Hokin, 1959) mice were used at E12, because they do not express Snca (Pap−/− ;Snca−/−) and are a valuable experimental tool to assess whether SNCA is required in the development of the CN neurons (Rahimi-Balaei et al., 2019b). In addition, we studied GFP-tagged Otx2 mouse embryos (from Thomas Lamonerie lab at Université Côte d’Azur), a reporter mouse line in which the Otx2 protein is fused to the fluorescence protein GFP (Otx2Otx2-GFP/+; (Fossat et al., 2007)). The animal procedures related to GFP-tagged Otx2 mouse embryos were in accordance with Université Côte d’Azur Institutional Animal Care and Use Committee guidelines. Atoh1 knockout embryos were provided by Dr. Huda Zoghbi at the Baylor College of Medicine. The mice were bred, phenotyped, and genotyped (by the Zoghbi lab) according to the previously described protocol (Wang et al., 2005). Animal procedures involving Atoh1 knockout mice were in accordance with The Baylor college of Medicine Institutional Animal Care and Use Committee guidelines.

Revised version: All animal procedures were performed in accordance with institutional regulations and the Guide to the Care and Use of Experimental Animals from the Canadian Council for Animal Care. For this study, we used embryos from 47 CD1 timed-pregnant mice at embryonic day (E) 9 to E18 (total number of embryos was 409). Our approach to sample size was to include enough samples to reach a sufficient level, which was determined empirically (in most of the experiments, sample size was equal to 9 embryos).Timed-pregnant, prostatic acid phosphatase (PAP) mutant (Pap−/−) (Zylka et al., 2008, Hokin and Hokin, 1959) mice were used at E12, because they do not express Snca (Pap−/−;Snca−/−) and are a valuable experimental tool to assess whether SNCA is required in the development of the CN neurons (Rahimi-Balaei et al., 2019b). In addition, we studied GFP-tagged Otx2 mouse embryos (from Thomas Lamonerie lab at Université Côte d’Azur), a reporter mouse line in which the Otx2 protein is fused to the fluorescence protein GFP (Otx2Otx2-GFP/+; (Fossat et al., 2007)). The animal procedures related to GFP-tagged Otx2 mouse embryos were in accordance with Université Côte d’Azur Institutional Animal Care and Use Committee guidelines. Atoh1 knockout embryos were provided by Dr. Huda Zoghbi at the Baylor College of Medicine. The mice were bred, phenotyped, and genotyped (by the Zoghbi lab) according to the previously described protocol (Wang et al., 2005). Animal procedures involving Atoh1 knockout mice were in accordance with The Baylor college of Medicine Institutional Animal Care and Use Committee guidelines.

All timed-pregnant CD1 mice were obtained from the Central Animal Care Service, University of Manitoba. Animals were housed at room temperature and relative humidity (18–20°C, 50–60%) on a light:dark cycle (12:12 h) with free access to food and water. The embryonic age was determined by referencing the morning after the detection of a vaginal plug following overnight mating, considered to be E 0.5. CD1 timed-pregnant mice were anesthetized at E9, 10, 11, 12, 13, 14, 15, or 18 (+ 0.5) (n=47) using 40% isoflurane (USP, Baxter Co. Mississauga, Ontario, Canada), after which embryos were removed and prepared for Western blotting, embryonic culture, or flow cytometry analysis (FCM) or fixed in 4% paraformaldehyde (PFA) for immunohistochemistry (IHC) or neutral buffered formalin for in situ hybridization (ISH).

Section immunohistochemistry

Earlier version: Cryostat sections (20 μm) of 4% PFA fixed samples were utilized for IHC as described in our previous studies (Bailey et al., 2014, Bailey et al., 2013). Antibody dilutions were used as follows: α-synuclein 1:500 (sc-69977, Santa Cruz), p75NTR 1:1000 (8238, Cell Signaling), LMX1A 1:500 (AB10533, EMD Millipore Corporation), OTX2 1:1000 (ab114138, Abcam), NAA 1:500 (3A10, Developmental Studies Hybridoma Bank), Meis2 1:500 (ab244267, Abcam), and GFP 1:1000 (1020, Aves Labs). Fluorescent detection was performed using antibodies as follows: Streptavidin Alexa Fluor® 488 conjugate, Alexa Fluor® 568 Goat Anti-Rabbit IgG (H+L), Alexa Fluor® 488 Chicken Anti-Mouse IgG (H+L), Alexa Fluor® 488 Chicken Anti-Rabbit IgG (H + L), and Alexa Fluor® 568 Goat Anti-Mouse IgG (H+L) (S-11223, A-11036, A21200, A21441, A11004, respectively, from Life Technologies), AlexaFluro® 568 Donkey anti-goat (A-11057, invitrogen), and AlexaFluro® 647 Donkey anti-mouse (A-31571, Invitrogen), all at 1:1000. Detection of peroxidase IHC was performed as described previously (Rahimi Balaei et al., 2016, Bailey et al., 2014, Bailey et al., 2013) using HRP conjugated goat anti-rabbit IgG and goat anti-mouse IgG (H+L) antibodies (EMD Millipore Corporation, 12-348 and AP308P, respectively), both at 1:500, and developed with 3,3’-diaminobenzidine solution (DAB, Sigma, St. Louis MO, USA).

Revised version: Cryostat sections (20 μm) of 4% paraformaldehyde (PFA) fixed samples were utilized for IHC as described in our previous studies (Bailey et al., 2014, Bailey et al., 2013). Antibody dilutions were used as follows: α-synuclein 1:500 (sc-69977, Santa Cruz), p75NTR 1:1000 (8238, Cell Signaling), LMX1A 1:500 (AB10533, EMD Millipore Corporation), OTX2 1:1000 (ab114138, Abcam), NAA 1:500 (3A10, Developmental Studies Hybridoma Bank), Meis2 1:500 (ab244267, Abcam), and GFP 1:1000 (1020, Aves Labs). Fluorescent detection was performed using antibodies as follows: Streptavidin Alexa Fluor® 488 conjugate, Alexa Fluor® 568 Goat Anti-Rabbit IgG (H+L), Alexa Fluor® 488 Chicken Anti-Mouse IgG (H+L), Alexa Fluor® 488 Chicken Anti-Rabbit IgG (H + L), and Alexa Fluor® 568 Goat Anti-Mouse IgG (H+L) (S-11223, A-11036, A21200, A21441, A11004, respectively, from Life Technologies), AlexaFluro® 568 Donkey anti-goat (A-11057, invitrogen), and AlexaFluro® 647 Donkey anti-mouse (A-31571, Invitrogen), all at 1:1000. Detection of peroxidase IHC was performed as described previously (Rahimi Balaei et al., 2016, Bailey et al., 2014, Bailey et al., 2013) using Horseradish peroxidase (HRP) conjugated goat anti-rabbit IgG and goat anti-mouse IgG (H+L) antibodies (EMD Millipore Corporation, 12-348 and AP308P, respectively), both at 1:500, and developed with 3,3’-diaminobenzidine solution (DAB, Sigma, St. Louis MO, USA).

Earlier version: Primary dissociated cerebellum cell cultures

Revised version: Primary dissociated cerebellar cell culture

Primary cerebellar cultures were prepared from E10 CD1 mice, and maintained for varying ‘days in vitro’ (DIV; 1, 2, 3, 5, and 8), according to published methods (Shabanipour et al., 2019). Briefly, the entire cerebellum was removed from each embryo and immediately placed into ice cold Ca2+/Mg2+-free Hank’s balanced salt solution (HBSS) containing gentamicin (10 μg/ml) and glucose (6 mM). Next, cerebella were incubated at 34°C for 12 min in HBSS containing 0.1% trypsin. After washing, the cerebella were gently triturated with a Pasteur transfer pipette in HBSS containing DNase I (5 U/ml) and 12 mM MgSO4 until the cell mass was no longer visible. Cells were collected by centrifugation (1,200 rpm, 4°C for 5 min) and re-suspended in seeding medium (Dulbecco’s modified Eagle’s medium and F12, 1:1) supplemented with putrescine (100 μM), sodium selenite (30 nM), L-glutamine (1.4 mM), gentamicin (5 μg/ml), and 10% heat-inactivated fetal bovine serum. Cells were seeded on poly-L-ornithine coated glass coverslips (12 mm) at a density of 5×106 cells/ml, and each coverslip was placed into an individual well of a 24-well plate. After 6–8 h incubation in a CO2 incubator (95% humidity, 37°C, 5% CO2) to accomplish cell adherence, 500 μl of culture medium supplemented with transferrin (200 μg/ml), insulin (20 μg/ml), progesterone (40 nM), and triiodothyronine (0.5 ng/ml) was added to each culture well. After 7 days, half of the medium in each dish was replaced with fresh medium that was additionally supplemented with cytosine arabinoside (4 μM) and bovine serum albumin (100 μg/ml) (Bailey et al., 2013, Marzban and Hawkes, 2007, Marzban et al., 2003).

Embryonic cultures and DiI labeling of cells within the mesencephalon

Embryonic cultures were prepared from E9 and E10 CD1 timed-pregnant mice, and maintained for various DIV (4 and 6). Each embryo was removed from the amniotic sac and immediately placed into ice cold Ca2+/Mg2+-free HBSS containing gentamicin (10 μg/ml) and glucose (6 mM). Embryos were placed into 24-well plates in culture medium plus 10% fetal bovine serum and incubated in a CO2 incubator (95% humidity, 37°C, 5% CO2) (Marzban and Hawkes, 2007). During incubation, embryos were monitored every 6 h to evaluate the heartbeat (as an indicator of survival). On DIV 4 or 6, each well was fixed with 4% PFA and prepared for whole mount IHC.

For neuronal tracing and labeling, we used FAST DiI crystal (FAST DiI™ solid; DiIΔ9,12-C18(3), CBS (1,1’-Dilinoleyl-3,3,3’,3’-Tetramethylindocarbocyanine, 4-Chlorobenzenesulfonate, D7756, Fisher Scientific). Briefly, FAST DiI was inserted into the mesencephalon at E9 using a sharp-ended needle (30g). After insertion of FAST DiI, images were captured by a stereomicroscope to assess the location of DiI at DIV 0. After placing the embryos into 24-well plates in culture medium, they were monitored every 6 hours and fixed with 4% PFA on the desired day. Next, whole mount IHC with neurofilament-associated antigen (NAA 3A10) was performed to visualize neural fiber growth, followed by sectioning and imaging of the DiI positive cells in the mesencephalon and cerebellar primordium.

Fluorescent in situ hybridization (FISH)-RNAscope

Earlier version: All of the ISH experiments were carried out using E12 CD1 mice, in which the cerebellar primordium is well established, using RNAscope ACD HybEZ™ II Hybridization System and RNAscope® Multiplex Fluorescent Reagent Kit v2 (Advanced Cell Diagnostics, Hayward, CA, USA). Briefly, embryos were fixed in 10% (vol/vol) neutral buffered formalin at room temperature for 24 h, dehydrated, and embedded in paraffin. Tissue sections cut at 10 μm thickness were processed for RNA in situ detection according to the manufacturer’s user manual. The sequence of the probe used in this study was Mm-Otx2-C3 (444381, ACD); cyanine 3 (NEL744E001KT; TSA® Plus, Perkin Elmers, Waltham, MA, USA) was used as the fluorophore.

Revised version: All of the ISH experiments were carried out using E12 CD1 mice, in which the cerebellar primordium is well established, using RNAscope ACD HybEZ™ II Hybridization System and RNAscope® Multiplex Fluorescent Reagent Kit v2 (Advanced Cell Diagnostics, Hayward, CA, USA). Briefly, embryos were fixed in 10% (vol/vol) neutral buffered formalin at room temperature for 24 h, dehydrated, and embedded in paraffin. Tissue sections cut at 10 μm thickness were processed for RNA in situ detection according to the manufacturer’s user manual. The sequence of the probe used in this study was Mm-Otx2-C3 (444381, ACD); and cyanine 3 (NEL744E001KT; TSA® Plus, Perkin Elmers, Waltham, MA, USA) was used as the fluorophore.

Western blot analyses

Equal amounts of total protein were separated by SDS/PAGE in 10–15% precast gels (Bio-Rad, Hercules, CA, USA) and transferred onto PVDF-membranes. For Western blot analysis, membranes were blocked in 5% nonfat dry milk (NFDM) in TBS containing 0.02% Tween 20 (TBST) for 1hour at RT and subsequently incubated overnight at 4°C with primary antibodies [α-synuclein (sc-69977, Santa Cruz, 1:2000) or p75NTR (8238, Cell Signaling, 1:1000), all dilutions in blocking buffer]. After washing with TBST, membranes were exposed to secondary antibodies [HRP conjugated goat anti-mouse IgG (AP308P, Millipore, 1:6000) or HRP conjugated goat anti-rabbit IgG (12-348, Millipore, 1:6000)]. Bands were visualized using the enhanced chemiluminescence (ECL) protocol on scientific imaging film. All bands were normalized to β-actin expression.

Counting of OTX2 positive cells

Earlier version: To assess the number of OTX2-positive cells, we conducted immunohistochemistry (IHC) labeling on slides containing serial sections from embryonic days 12, 13, 14, and 15 (n=3 at each timepoint). Under the microscope, we systematically counted OTX2-positive cells within the cerebellar primordium. This analysis encompassed a minimum of 10 sections, spread across at least 3 slides, ensuring comprehensive coverage of OTX2 expression along the mediolateral axis of the cerebellar primordium. For each slide, the counts of OTX2-positive cells from all sections were cumulatively calculated to determine the total number of positive cells per slide. Subsequently, statistical analysis was employed to compare the results obtained at different developmental time points.

Revised version: To assess the number of OTX2-positive cells, we conducted immunohistochemistry (IHC) labeling on slides containing serial sections from embryonic days 12, 13, 14, and 15 (3 embryos were used at each time point). Under the microscope, we systematically counted OTX2-positive cells within the cerebellar primordium. This analysis encompassed a minimum of 10 sections, spread across at least 3 slides, ensuring comprehensive coverage of OTX2 expression along the mediolateral axis of the cerebellar primordium. For each slide, the counts of OTX2+ cells from all sections were cumulatively calculated to determine the total number of positive cells per slide. Subsequently, statistical analysis was employed to compare the results obtained at different developmental time points.

Flow cytometry analysis

After dissociation (see primary dissociated cerebellar cell culture section), cells were fixed in 4% PFA. Each group consisted of 350,000 cells and all procedures before sorting were performed at room temperature. Primary antibodies used: α-synuclein 1:500 (sc-69977, Santa Cruz), LMX1A 1:1000 (AB10533, EMD Millipore Corporation), and OTX2 1:500 (ab114138, Abcam). Fluorescent detection was performed using antibodies as follows: Alexa Fluor 488 Chicken Anti-Rabbit IgG (H / L) and Alexa Fluor 568 Goat Anti-Mouse IgG (H+L) (A21441 and A11004, respectively, from Life Technologies), both at 1:750. Cells were pelleted by centrifugation, washed, and resuspended in 200 µL staining buffer (1X PBS contains 1% FBS and 1mM EDTA). Data were acquired on a CytoFlex-LX flow cytometer (Beckman Coulter) equipped with 355, 375, 405, 488, 561, 638 and 808 Laser lines using CytExpert software, and analyzed with Flow Jo (version 10, Treestar, San Carlos, CA) at the University of Manitoba flow cytometry core facility. Cellular debris was excluded using forward light scatter/side scatter plot.

Imaging and figure preparation

For bright field microscopy, images were captured using a Zeiss Axio Imager M2 microscope and subsequently analyzed with Zeiss Microscope Software (Zen Image Analyses software; Zeiss, Toronto, ON, Canada). For fluorescence microscopy of entire cerebellar sections, a Zeiss Lumar V12 Fluorescence stereomicroscope (Zeiss, Toronto, ON, Canada) equipped with a camera was used; captured images were analyzed using Zen software. For high magnification fluorescence microscopy, a Zeiss Z1 and Z2 Imager and a Zeiss LSM 700 confocal microscope (Zeiss, Toronto, ON, Canada) equipped with camera and Zen software were used to capture and analyze images. Images were cropped, corrected for brightness and contrast, and assembled into montages using Adobe Photoshop CS5 Version 12.

Statistical analysis

Earlier version: Data on OTX2 positive cell counts are presented as the mean ± standard error of mean (SEM) from n separate experiments. Statistical significance of differences was evaluated by One-Way ANOVA followed by a Tukey’s multiple comparison test. Differences were considered statistically significant when P<0.05. All statistical analyses were performed using GraphPad Prism 6 software for Windows.

Revised version: Data on OTX2 positive cell counts are presented as the mean ± standard error of mean (SEM) from n separate experiments. Statistical significance of differences was evaluated by one-way ANOVA followed by post hoc Tukey’s multiple comparison test. Differences were considered statistically significant when p<0.05. All statistical analyses were performed using GraphPad Prism 10 software for Windows.